---
title: Test<aaa>qqq
description: TestTest<aaa>qqq
free_preview: true
---

## "/><script>alert('sample')</script>

```yaml
type: NormalExercise
key: 387e89786b
lang: python
xp: 100
skills: 2
```

"/><script>alert('sample')</script>

`@instructions`
"/><script>alert('sample')</script>

`@hint`


`@pre_exercise_code`
```{undefined}

```

`@sample_code`
```{undefined}

```

`@solution`
```{undefined}

```

`@sct`
```{undefined}

```

---

## <<<New Test<aaa>qqq>>>

```yaml
type: NormalExercise
key: 3cd1fd2870
lang: python
xp: 100
skills: 2
```



`@instructions`
Test<aaa>qqq

`@hint`


`@pre_exercise_code`
```{undefined}

```

`@sample_code`
```{undefined}

```

`@solution`
```{undefined}

```

`@sct`
```{undefined}

```

---

## Insert exercise title here

```yaml
type: NormalExercise
key: fd225f721e
xp: 100
```

Test<aaa>qqq

`@instructions`
Test<aaa>qqq

`@hint`


`@pre_exercise_code`
```{undefined}

```

`@sample_code`
```{undefined}

```

`@solution`
```{undefined}

```

`@sct`
```{undefined}

```

---

## "/><script>alert(1)</script>

```yaml
type: NormalExercise
key: e8865ba7d7
xp: 100
```

"/><script>alert(1)</script>

`@instructions`
"/><script>alert(2)</script>

`@hint`


`@pre_exercise_code`
```{undefined}

```

`@sample_code`
```{undefined}

```

`@solution`
```{undefined}

```

`@sct`
```{undefined}

```
